<a href="https://colab.research.google.com/github/Riab22209/CE880-Lab-coursework/blob/main/week1/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Welcome to CE880
### This is your week-1 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function for Pythagorean theorem, or Pythagoras' theorem. Pythagorean is a fundamental relation in Euclidean geometry among the three sides of a right triangle. It states that the area of the square whose side is the hypotenuse (the side opposite the right angle) is equal to the sum of the areas of the squares on the other two sides. This theorem can be written as an equation relating the lengths of the legs a, b and the hypotenuse c, often called the Pythagorean equation:

$a^2 + b^2 = c^2$

Please write a function (code)to calculate the length of an unknown side of a right triangle, given two sides are known.

<!-- grade_id: myPythagorean_ans, points: 0, grade: false, solution: true -->
```python
import math as mt
def myPythagorean(a,b):
    """write a function (code)to calculate the length of an unknown side 
    of a right triangle, given two sides are known.
    """
    # YOUR CODE HERE
    """write a function (code)to calculate the length of an unknown side 
    of a right triangle, given two sides are known.
    """
    if (a==0 or b ==0):
    
        return 0
    
    else:
      
          c = mt.sqrt(a**2 + b**2)
          return c 

print("The unknown side of the right angled triangle is : ")
myPythagorean(10,14)
#raise NotImplementedError()

```

Output:
```
The unknown side of the right angled triangle is : 
17.204650534085253
```

```python
myPythagorean(10,14)
```

Output:
```
17.204650534085253
```

<!-- grade_id: myPythagorean_test, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert myPythagorean(3,2) == 3.605551275463989
assert myPythagorean(7,3) == 7.615773105863909
assert myPythagorean(0,2) == 0
```

---

## Question 2:

A `trapezoid` is a quadrilateral with two sides parallel. The `trapezoid` is equivalent to the British definition of the `trapezium`. An isosceles `trapezoid` is a `trapezoid` in which the base angles are equal so

<img src="https://www.w3resource.com/w3r_images/python-math-image-exercise-3.svg"
     alt="Markdown Monster icon"
     style="float: left; margin-right: 10px;" />
     
Please write a function (code)to calculate the area of trapezoid.

<!-- grade_id: areaofTrap_ans, points: 0, grade: false, solution: true -->
```python
def areaofTrap(base1, base2, height):
    """Compute the sum of the squares of numbers from 1 to n."""
    # YOUR CODE HERE
    A = ((base1+base2)*height)/2
    return A
print("The area of the Trapezoid is: ")   
areaofTrap(5, 3, 6)
#raise NotImplementedError()
```

Output:
```
The area of the Trapezoid is: 
24.0
```

```python
areaofTrap(5, 3, 6)
```

Output:
```
24.0
```

<!-- grade_id: areaofTrap_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert areaofTrap(10, 10, 20) == 200.0
assert areaofTrap(5, 10, 20) == 150.0
assert areaofTrap(5, 3, 5) == 20.0
```

---
## Question 3: 

 Write a Python program to reverse a string. Example: Sample String: `1234xyz` Expected Output: `zyx4321`

<!-- grade_id: reversestring_ans, points: 0, grade: false, solution: true -->
```python
def reverse_string(str1):
    """Write your code here to reverse the string"""
    rev_string = ""
    for i in str1:
        rev_string = i + rev_string
    return rev_string
print("The original string is RIA")   
print("The reversed string is : ")    
reverse_string("RIA")
    # YOUR CODE HERE
    #raise NotImplementedError()
```

Output:
```
The original string is RIA
The reversed string is : 
'AIR'
```

```python
reverse_string('ChatGPT')
```

Output:
```
'TPGtahC'
```

<!-- grade_id: reversestring_test, points: 1, grade: true, solution: false -->
```python
"""Check that reverse_string returns the correct answer for various inputs."""
assert reverse_string('Humanitarian') == 'nairatinamuH'
assert reverse_string('ChatGPT') == 'TPGtahC'
assert reverse_string('Data Science') == 'ecneicS ataD'
assert reverse_string('CE880') == '088EC'
```

## Question 4:
Write a Python program to print the even numbers from a given list and in ascending order.

Sample List : `[1, 2, 3, 4, 5, 6, 7, 8, 9]` Expected Result : `[2, 4, 6, 8]`

<!-- grade_id: even_num_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def even_num(lst_1):
    """Write a Python program to print the even numbers from a given list and in ascending order."""
    # YOUR CODE HERE
      # checking condition
    even_nos = [num for num in lst_1 if num % 2 == 0]
    even_nos=list(set(even_nos)) # remove duplicate values from list
    even_nos.sort() # sort the list in ascending order
    return even_nos
print("List of even numbers in ascending order:")  
even_num([1,2,3,4,5,6,7,8,9])         
    #raise NotImplementedError()
```

Output:
```
List of even numbers in ascending order:
[2, 4, 6, 8]
```

```python
even_num([2, 3, 14, 17, 8, 19, 11, 14])
```

Output:
```
[2, 8, 14]
```

<!-- grade_id: odd_num_test, points: 1, grade: true, solution: false -->
```python
assert even_num([1, 2, 3, 4, 6, 7, 8, 9]) == [2, 4, 6, 8]
assert even_num([2, 3, 4, 7, 8, 9, 11, 14]) == [2, 4, 8, 14]
assert even_num([2, 3, 14, 17, 8, 19, 11, 14]) == [2, 8, 14]
```

## Question 5:
Write a Python function to create and print a list where the values are square of numbers between `1` and `n`, where `n` is user defined number: Example `sqr_values(4) == [1, 4, 9, 16]`

<!-- grade_id: sqr_values_ans, points: 0, grade: false, solution: true -->
```python
def sqr_values(n):
    """Write a Python program to print the square of numbers between 1 and `n`""" 
    # YOUR CODE HERE
    return list(map(lambda k:k**2, range(1,n+1)))
sqr_values(2)    
    #raise NotImplementedError()
```

Output:
```
[1, 4]
```

```python
sqr_values(4)
```

Output:
```
[1, 4, 9, 16]
```

<!-- grade_id: sqr_values_test, points: 1, grade: true, solution: false -->
```python
"""Check that sqr_values returns the correct answer for various inputs."""
assert sqr_values(4) == [1, 4, 9, 16]
assert sqr_values(8) == [1, 4, 9, 16, 25, 36, 49, 64]
assert sqr_values(9) == [1, 4, 9, 16, 25, 36, 49, 64, 81]
assert sqr_values(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

## Question 6: 

Write a Python program to find all the duplicate characters which are similar to each other and join them.


<!-- grade_id: comb_ans, points: 0, grade: false, solution: true -->
```python
from ast import JoinedStr
def findDupChar(input):
    """Write a Python program to find all the duplicate characters which are similar to each other."""
    # YOUR CODE HERE
    x = filter(lambda x: input.count(x) >= 2, input)
    x = set(x)
    print(x)
    #print(' '.join(x))
   
 
    #return duplicates
findDupChar("Happyy") 
    #raise NotImplementedError()
```

Output:
```
{'y', 'p'}
```

```python
findDupChar('Essex University')
```

Output:
```
s i e
```

<!-- grade_id: comb_test, points: 1, grade: true, solution: false -->
```python
assert findDupChar('Happyy') == 'p y'
assert findDupChar('Essex University') == 's e i'
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-108-0010d4b24671> in <module>
----> 1 assert findDupChar('Happyy') == 'p y'
      2 assert findDupChar('Essex University') == 's e i'

AssertionError:
```

```python

```